---
title: 'Frequency Tables'
description: 'In general a large data set is called a raw data set.  As a first step in making sense of an unorganized data set,  it must be organized into a frequency table.  The R has many packages to accomplish the task of  building a frequency table.  Frequency tables are mainly two types: Ungrouped type and Grouped type.  The main aspects in a frequency table are frequency, cumulative frequency,relative frequency and relative cumulative frequency .   Once we produce such tables, we need to learn how to interpret the results.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

Assume that there are 100 students in a class who all took a quiz that is worth 40 points. Using R generate 100 random integers and construct an ungrouped frequency table.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
